---
title: "Grading Code with gradethis"
subtitle: "satRday 2020"
author: "Daniel Chen @chendaniely"
institute: "Virginia Tech"
date: "2020-03-28"
output:
  xaringan::moon_reader:
    css: ["footer-header.css", "slide-content.css", "default"]
    lib_dir: libs
    nature:
      highlightStyle: github
      highlightLines: true
      countIncrementalSlides: false
      ratio: "16:9"
editor_options: 
  chunk_output_type: console
---

class: inverse

<div class="my-footer"><span>xaringan power    
&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;
&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;&emsp;
yolo</span></div> 

<center>
<div class="hello">hi!</div>
</center>

---

class: inverse

# I'm Daniel
.pull-left[
<img src='./figs/daniel_square-800x800.jpg'></img>
]

.pull-right[
- PhD Student: Virginia Tech
  - Data Science education & pedagogy
  - Medical, Biomedical, Health Sciences
  - Advisor: Anne Brown, PhD
      - https://www.databridge.dev/
      - https://bevanbrownlab.com/
-  Inten at RStudio
  - [`gradethis`](https://github.com/rstudio-education/gradethis)
  - Code grader for [`learnr`](https://github.com/rstudio/learnr) documents
- Author:
<center>
<img src='./figs/book.jpg' style="width:200px;"></img>
</center>
]

---

# RStudio Internship

https://rstudio-education.github.io/gradethis/

- Worked on `gradethis` for my [RStudio 2019 Internship](https://blog.rstudio.com/2019/03/25/summer-interns-2019/)
  - Barret Schloerke, PhD
  - Garrett Grolemund, PhD

- I am a user of R, not a developer.

- Wrote about my experiences:
  - [My time as an intern](https://daniel.rbind.io/2020/01/29/my-time-as-an-rstudio-intern/)
  - [Internship week 1](https://daniel.rbind.io/2019/06/10/and-were-off-rstudio-internship-week-1-complete./)
  - [Internship week 2](https://daniel.rbind.io/2019/06/18/rstudio-internship-week-2/)
  - [Moved to blogdown](https://daniel.rbind.io/2019/07/23/new-website-a-la-blogdown/)
  - [Git workflow](https://daniel.rbind.io/2019/08/27/git-squash-and-merge-workflow/)
  - [Rstudio education blog post](https://education.rstudio.com/blog/2020/02/gestalt-internship/)

---

# Education + Shiny Team

.pull-left[
Education Team

![](./figs/00-education_team.png)
]

.pull-right[
Shiny Team

![](./figs/00-shiny_team.png)
]

---

# Join all the slack channels!

- What? A book about shiny?

- Yes! https://mastering-shiny.org/
- No ):

![](./figs/shiny_book-join_leave.png)



---

# Learnr interactive tutorials

https://rstudio.github.io/learnr/


1. Narrative, figures, illustrations, and equations.
2. Code exercises (R code chunks that users can edit and execute directly).
3. Quiz questions.
4. Videos (supported services include YouTube and Vimeo).
5. Interactive Shiny components.


---

# Creating a `learnr` tutorial

```yaml
---
title: "Hello, Tutorial!"
output: learnr::tutorial
runtime: shiny_prerendered
---
```

````markdown
`r ''````{r setup, include=FALSE}
library(learnr)
```

This code computes the answer to one plus one, change it so it computes two plus two:

`r ''````{r addition, exercise=TRUE}
1 + 1
```
````

---

# Learnr output

<center>
<img src="./figs/learnr-output.png" />
</center>

---

# Sortable

- More `learnr` examples: https://rstudio.github.io/learnr/examples.html
  - `sortable`: https://andrie-de-vries.shinyapps.io/sortable_tutorial_question_rank/

<center>
<img src="./figs/sortable-example.gif" />
</center>

---

# Sortable -> Parsons

- Parson problems are coding exercises designed to **reduce cognitive load** on the learner.
  - Instead of writing the code, they order code blocks in the correct order
  - Reduces syntax errors
    - Learn the "what" concepts instead of the "how" syntax

https://github.com/rstudio/parsons

---

# Parsons

<center>
<img src="./figs/parsons-example.gif" />
</center>

---

# Grading the questions

We can create questions, but how do we give feedback?

At the very least how to we tell the student the solution they provided was "wrong"?

---

# Look at the result

- Only check if the solutions match
- As long as the correct answer is returned, it doesn't matter how the student got the answer
  - Can't check if you are trying to teach a specific function

E.g., `3 + 3 = 6 = 2 * 3`

---

# Look at the code

- Essentially compares the actual code (e.g., the text)
- It's very strict
  - `3 + 3` != `2 * 3`
  - `apply(df, 2, class)` != `apply(df, MARGIN = 2, FUN = class)`
  - `mean(1:5)` != `mean(1:5, na.rm=FALSE)`

---

# Look at the code's AST

- The Abstract Syntax Tree (AST) is a graphical (i.e., tree) representation of
your code.
- It shows the order of function calls and their arguments to be executed.

.pull-left[
```{r}
lobstr::ast(
  log(exp(3))
)
```

```{r}
lobstr::ast(
  log(exp(2))
)
```
]

.pull-right[
```{r}
lobstr::ast(
  log(log(2, base = 10))
)
```
]

---

# Gradethis

https://rstudio-education.github.io/gradethis/

- Auto grading system that can be used to grade `learnr` exercises.
  - Check the answer
  - Check the code (using the AST)

- In `learnr`, each exercise is completely independent from the others

- What makes it unique is its ability to give formative feedback to the student.

- `leanr` but why not `grader`?
  - The package was originally named `grader`
  - Someone from UVA took the name on CRAN during the first few weeks of the internship.
  - https://github.com/rstudio-education/gradethis/issues/18

---

# Formative feedback

<center>
<img src="./figs/gradethis-example.png" />
</center>

---

# `learnr` + `gradethis`

1. base chunk: what the student sees in the exercise
2. `-check` chunk: how should the student's code be graded
    1. solution: are you just checking the final result?
        - `gradethis::grade_result()`
        - Provide a set of conditions to check, returns result on *first match* (i.e., order of conditions matter)
    2. code: are you checking to see the code itself is correct (AST)
        - `gradethis::grade_code()`
        - `-solution` chunk: The instructor's solution code that will be used to compare to the student's code
    3. unittests: are you testing a function that needs to pass a bunch of checks?
        - `gradethis::grade_conditions()`
3. `-hint` chunks (optional)

---

# `learnr`

Set the checking function for `learnr`

````markdown
`r ''````{r setup}
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr)
```

We can scaffold the exercise to reduce cognitive load.
````markdown
`r ''````{r addition, exercise=TRUE}
log(_____(_____))
```
````

And want the student to enter

````markdown
`r ''````{r addition, exercise=TRUE}
log(exp(3))
```
````

---

# `learnr` + `gradethis::grade_result()`

````markdown
`r ''````{r addition, exercise=TRUE}
log(_____(_____))
```

`r ''````{r addition-hint-1}
# hint text
"Exponentiate with the `exp` function."
```

`r ''````{r addition-check}
# check result
gradethis::grade_result(
  # pass_if fail_if order does matter. Maybe put pass_if conditions last?

  #.result is the last value returned by the student
  gradethis::pass_if(~ identical(.result, 3), "YAY!"),
  gradethis::fail_if(function(x){identical(x, 4)}, "4 is an incorrect input.")
  gradethis::fail_if(2, "2 was the wrong number to use here.")
)
```
````

---

# `learnr` + `gradethis::grade_code()`

````markdown
`r ''````{r addition, exercise=TRUE}
log(_____(_____))
```

`r ''````{r addition-hint-1}
# hint text
"Exponentiate with the `exp` function."
```

`r ''````{r addition-solution}
# solution code
log(exp(3))
```

`r ''````{r addition-check}
# check code
gradethis::grade_code()
```
````

---

# `learnr` + `gradethis::grade_conditions()`

````markdown
`r ''````{r grade_conditions, exercise = TRUE}
# student code
add_1 <- function(x) {x + 1}
```

`r ''````{r grade_conditions-hint-1}
# hint text
"Function should take a parameter and add 1 to it"
```

`r ''````{r grade_conditions-check}
gradethis::grade_conditions(
  gradethis::pass_if(~ .result(3) == 4),
  gradethis::pass_if(~ .result(-1) == 0),
  gradethis::fail_if(~ .result(10) == 11)
)
```
````

---

# Chunk names matter

.pull-left[
````markdown
`r ''````{r addition, exercise=TRUE}
log(_____(_____))
```

`r ''````{r addition-hint-1}
# hint text
"Exponentiate with the `exp` function."
```

`r ''````{r addition-solution}
# solution code
log(exp(3))
```

`r ''````{r addition-check}
# check code
gradethis::grade_code()
```
````
]

.pull-right[
- The hints/soluction/check blocks use the same base name for association
- `knitr` chunks need to be unique
]

---

# `gradethis` chunk addins

.pull-left[
- Creates random chunk names

<img src="./figs/gradethis-addon.png" width='100%' />
]

.pull-right[
````markdown
`r ''````{r qhdcovmzxfhuycxw, exercise = TRUE}
# student code
____
```

`r ''````{r qhdcovmzxfhuycxw-hint-1}
# hint text
""
```

`r ''````{r qhdcovmzxfhuycxw-hint-2}
# hint text
""
```

`r ''````{r qhdcovmzxfhuycxw-solution}
# solution code

```

`r ''````{r qhdcovmzxfhuycxw-check}
# check code
gradethis::grade_code()
```
````
]
---

# Be careful with `==`

- https://daniel.rbind.io/2019/08/06/inconsistencies-with-in-r/
  - "inconsistency", "improper documentation", "bug"

```{r}
# only difference is the last parameter
u <- quote(f(x123456789012345678901234567890123456789012345678901234567890, 1))

s <- quote(f(x123456789012345678901234567890123456789012345678901234567890, 2))
```

```{r}
u == s
```

```{r}
identical(u, s)
```

- Take away: use `idential` or `all.equal` instead of `==` when doing (non-vectorized) comparisons.
  - Vignette in PR: https://github.com/rstudio-education/gradethis/pull/102

---

# Dev guide: `learnr` vs `gradethis`

- `learnr` is responsible for
  - Expression capturing
  - Data reporting
    - How many questions did my student get "correct"?
    - How much time did my students take for each question or for the entire document?

- `gradethis` is responsible for:
  - Checking the student solution
  - Correct/Incorrect message + praise/encouragement

---

# Dev guide: `learnr` checker

- `learnr` is a complex piece of software that needs a deep understanding of `knitr`, `rmarkdown`, and `shiny`.

- All of the exercise components in the `learnr` document are passed into the `exercise.checker`.
  - This is an entry point if you want to write your own `learnr` checker

- To enable exercise checking in your learnr tutorial:

````markdown
`r ''````{r setup}
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr)
```
````

- The `grade_learnr` function:
  - The function `learnr` passes all the chunk code into
  - Passes into the corresponding `gradethis::grade_` function

---

# Standalone code grader: result

```{r}
library(gradethis)

student_solution <- 5
```

```{r}
graded_result <- gradethis::grade_result(
  pass_if(5, "You also got a value of 5!"),
  learnr_args = list(last_value = student_solution)
)

graded_result
```

---

# Standalone code grader: code

```{r}
student_code  <- quote(sqrt(exp(3)))
solution_code <- quote(sqrt(log(2)))
```

```{r}
graded_code <- gradethis::grade_code(grader_args = list(
  user_quo = student_code,
  solution_quo = solution_code
  )
)

graded_code
```
---
class: inverse

# User -> Dev
---

# Internship: `eval_tidy` example

```{r}
form <- ~ x ^ 2 # capture an expression + environment
str(form)
```

```{r}
form[[2]] # just get the expression
```

```{r}
rlang::eval_tidy(
  expr = form[[2]], # evaluate the expression with data + environment
  data = list(x = 15),
  env = .GlobalEnv
)
```

---

# Internship: `eval_tidy` in practice

https://github.com/rstudio-education/gradethis/blob/master/R/evaluate_condition.R

During `gradethis::grade_result`:

```r
switch(condition$type,
           "formula" = evaluate_condi_formula(condition$x, learnr_args$last_value, learnr_args$envir_prep), # nolint
           "function" = evaluate_condi_function(condition$x, learnr_args$last_value),
           "value" = evaluate_condi_value(condition$x, learnr_args$last_value)
         )
```

```r
evaluate_condi_formula <- function(formula, user_answer, env) {
  form_result <- rlang::eval_tidy(
    formula[[2]],
    data = list(.result = user_answer, . = user_answer),
    env = env
  )
  return(form_result)
}
```

---

# Internship: `lapply` good

```{r}
vec <- 1:10
```

```{r}
lapply(vec, log, base = 10)
```

---

# Internship: `sapply` bad

```{r}
# simplify this value into a vector, if you can't just give me a list
sapply(vec, log, base = 10)
```

---

# Internship: `vapply` better

```{r}
# I want a numerical vector back where each element only contains 1 element
vapply(vec, FUN = log, FUN.VALUE = numeric(1), base = 10)
```

---

class: center, middle, inverse


# Thanks!

.pull-left[
<img src="./figs/interns.gif" />
]

.pull-right[
@chendaniely

Slides: https://github.com/chendaniely/
]
